Dieser Kurs basiert aus dem Einführungskurs für das [physikalische Anfängerpraktikum der Universität Heidelberg](https://www.physi.uni-heidelberg.de/Einrichtungen/AP/Python.php)

---
# Herzlich willkommen zum Einblick in die Signalverarbeitung mit Python!

In diesem Kurs bekommen Sie einen Einblick in die Programmiersprache Python sowie in einige Grundlagen der Signalverarbeitung. Python ist im Gegensatz zu C eine Interpretersprache. Interpreter erzeugen keinen Maschinencode, den der Computer direkt ausführt. Stattdessen werden die Python-Programme durch einen Interpreter in Echtzeit Zeile für Zeile übersetzt und ausgeführt. Interpretierte Programme sind daher grundsätzlich langsamer als kompilierte, bieten dafür dem Programmierer mehr Möglichkeiten.

Sie lernen in diesem Kurs, wie in Python mit Variablen, Typen, Listen, Dictionaries, Kontrollstrukturen und Arrays umgegangen sowie Plots erstellt werden.

Der Kurs basiert auf einer speziellen Entwicklungsumgebung, den *Jupyter-Notebooks*.

# Jupyter-Notebooks

Programmiercode ist erst einmal Text. Ein Python-Skript wird normalerweise als Quelltext in `.py` Dateien gespeichert, doch für wissenschaftliche Anwendungen sind solche Skripte häufig nicht anschaulich genug. 

Stattdessen verwenden wir in diesem Kurs [_Jupyter Notebooks_](http://jupyter.org), in denen wir nicht nur Code ausführen sondern auch Text, $\LaTeX$-Gleichungen, Bilder und mehr hinzufügen können. In diesem Kurs werden Sie den Umgang mit Notebooks noch kennenlernen, schließlich sind diese Kursmaterialien selbst Jupyter Notebooks!

Das Jupyter Notebook muss dabei immer auf einen _Kernel_ zurückgreifen können, also eine Python Installation auf dem System. Falls der Kurs nicht auf den Rechnern der HTW Berlin ausgeführt wird, müssen Python und Jupyter auf dem Zielrechner installiert werden. Unter Linux ist die pip-Methode in der Regel die geeignetere, aber auch die unten beschriebene Installation über Anaconda führt zum Ziel.

## Was ist ein Jupyter-Notebook?

Ein Jupyter-Notebook ist ein webbasiertes Dokument, welches aus Zellen besteht. Es gibt drei Arten von Zellen:
- Markdown-Zellen. Sie enthalten formatierte Texte, Kommentare usw. Die Texte werden in einem Markdown-Format beschrieben. Dieser Text steht in einer Markdown-Zelle.
- Codezellen (Eingabezellen). Sie enthalten Python-Code, der ausgeführt werden soll
- Ausgabezellen. Sie enthalten die Ausgaben von dem Python-Code in der vorhergehenden Eingabezelle.

Markdown-Zellen und Codezellen haben zwei Modi:
- Editiermodus: Hier kann der Python- bzw. Markdown-Code eingegeben werden
- Rendermodus: Hier ist der Code bzw. Text nur sichtbar. Im Fall von Markdownzellen ist der Text in einer gerenderten Form lesbar.

Durch einen Doppelklick können Sie eine Zelle in den Editiermodus bringen. (Das funktioniert auch bei den Texten in Markdownzellen).
Durch `Shift+Return` schalten Sie eine Zelle wieder in den Rendermodus um. In Eingabezellen (Codezellen) wird durch Shift+Return der Python-Quelltext ausgeführt und das Ergebnis in die nachfolgende Ausgabe geschrieben.

Probieren Sie es aus:

```python
x = 1.5
# Diese Zelle auswählen und mit `<SHIFT> + <ENTER>` ausführen.
# Der letzte Rückgabewert der Zelle wird dann unten ausgegeben,
# hier also der Wert von `x`:
x
```

Mehr Informationen über Jupyter-Notebooks gibt es beispielsweise unter https://jupyter-tutorial.readthedocs.io/de/latest/. Jupyter-Notebooks sind ein Standardwerkzeug in den Bereichen Data Science und maschinelles Lernen.

### Jupyter-Notebooks in den Übungen

Das Format der Jupyter Notebooks eignet sich hervorragend für die Versuche im Rahmen der Übungen. Erstellen Sie einfach einen Ordner für jeden Versuch und erstellen Sie darin ein Jupyter Notebook. Speichern Sie die Messdaten und Bilder ebenfalls in diesem Ordner, um sie im Notebook einfach laden zu können.

# Grundlagen der Programmierung in Python

Programmieren in Python besteht aus einer Aneinanderreihung von Befehlen, die der Interpreter ausführt. Ein einfacher Befehl ist die Zuweisung eines Wertes zu einer **Variablen**:

```python
x = 1
# Diese Zelle auswählen und mit `<SHIFT> + <ENTER>` ausführen.
# Der letzte Rückgabewert der Zelle wird dann unten ausgegeben,
# hier also der Wert von `x`:
x
```

Output:
```
1
```

Für den weiteren Verlauf des Kurses sind folgende Eingaben notwendig. Führen Sie diese mit `<SHIFT> + <ENTER>` aus. Wenn Sie dieses Notebook neu starten, müssen Sie diese Zeilen auf jeden Fall ausführen, damit es nicht zu unvorhergesehen Fehlermeldungen kommt:

```python
%matplotlib inline
import numpy as N
from matplotlib import pyplot as P
import scipy.signal as SI

def assert_equal(x, y, s):
    if (x != y):
        raise ValueError(s)
```

> **Hinweis:** Dieses Kursmaterial besteht aus interaktiven [Jupyter Notebooks](http://jupyter.org). Jede _Zelle_, wie die obige, können Sie per Klick oder mit den Pfeiltasten auswählen und mit `<SHIFT> + <ENTER>` ausführen. Probieren Sie es mal mit der Zelle oben!
> 
> **Gehen Sie, um dem Kurs zu folgen, _alle_ Zellen nacheinander durch und führen Sie sie aus.**
>
> Wie hier demonstriert gibt das Jupyter Notebook immer den letzten Rückgabewert einer Zelle aus. Wir können auch die `print()` Funktion verwenden, um Ausgaben zu generieren:

```python
print(x)
```

Der **Wert** der Variable `x` ist nun der Integer `1`. Werte haben immer einen **Typ** wie

- `int` für ganze Zahlen, z.B. `1`, `42`, `-10`,
- `float` für Fließkommazahlen, z.B. `0.5`, `3.14`, `1e10`,
- `str` für Zeichenketten (_Strings_), z.B. `"Hello World!"`,
- `boolean` für Wahrheitswerte (_Booleans_), also `True` und `False`

> Die Typen der Python Standardlibrary findet ihr in der [Dokumentation](https://docs.python.org/3/library/stdtypes.html).

Da Variablen in Python nicht typisiert sind, können der gleichen Variable nacheinander Werte verschiedener Typen zugewiesen werden:

```python
type(x)
```

```python
x = 0.5
type(x)
```

Weiterhin können Werte in einen anderen Typ konvertiert werden:

```python
x = int(0.5) # Bei der Konvertierung zu int wird immer abgerundet!
x
```

Die grundlegenden Rechenoperationen `+`, `-`, `*`, `/` und `**` sind ebenfalls in Python verfügbar und verhalten sich, wie man es erwartet:

```python
1 + 3
```

```python
3 / 2
```

```python
3**2
```


## Aufgabe 1 - Werte und Typen

Der Operator zur Potenzierung ist in Python `**`. Weisen Sie einer Variablen `y` folgende Werte zu und lassen Sie dann ihren Wert und Typ ausgeben. Stimmen Wert und Typ mit Ihren Erwartungen überein?

a) $4^3$

> **Hinweis:** Dies ist eine Übungsaufgabe. Verwenden Sie die Zelle unten, um sie zu lösen. **Entfernen** Sie dazu den Platzhalter-Code
>
> ```python
> # IHRE LÖSUNG HIER
> ```
>
> und schreiben Sie stattdessen den Code zur Lösung der Aufgabe.

<!-- grade_id: 101-1a-sol, points: 0, grade: false, solution: true -->
```python
# IHRE LÖSUNG HIER
```

> **Hinweis:** Auf die Lösung jeder Aufgabe folgt eine Test-Zelle wie die folgende, mit der Sie Ihre Lösung überprüfen können. Führen Sie einfach diese Zelle aus. Wenn etwas an Ihrer Lösung nicht stimmt, erhalten Sie eine Fehlermeldung mit Hinweisen.

<!-- grade_id: 101-1a-test, points: 1, grade: true, solution: false -->
```python
try:
    y
except NameError:
    raise NameError("Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_equal(y, 64, "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung.")
print("👍 Sehr gut!")
```

b) $2+3.4^2$

<!-- grade_id: 101-1b-sol, points: 0, grade: false, solution: true -->
```python
# IHRE LÖSUNG HIER
```

<!-- grade_id: 101-1b-test, points: 1, grade: true, solution: false -->
```python
try:
    y
except NameError:
    raise NameError("Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_equal(y, 2 + 3.4**2, "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung.")
print("Ok, du hast es verstanden.")
```

## Listen

Ein weitere Variablentyp in Python ist die **Liste**. Eine Liste speichert eine Reihe _beliebiger_ Werte und kann dabei auch Werte unterschiedlicher Typen enthalten. Die Elemente einer Liste können verändert werden, während ein Tupel unveränderlich ist. Wir können Listen erstellen, indem wir die Elemente durch Kommata getrennt in eckigen Klammern `[` und `]` angeben:

```python
l = [ 4, 0.5, "Alice" ]
l
```

```python
t = ( "Bob", True )
t
```

Wir können die Anzahl der Elemente bestimmen und auf einzelne Elemente über ihren Index zugreifen. **Indices in Python beginnen wie in C mit Null:**

```python
len(l), len(t)
```

```python
l[0] # Indexierung beginnt in Python mit dem Index 0
```

```python
l[1]
```

```python
l[2]
```

```python
t[0]
```

```python
t[1]
```

## Negative Indizes

Sie können mit der _Slicing_ Syntax auf Teile einer Reihe zugreifen:

```python
slice = list[start:stop:step]
```

Dabei bezeichnet `start` den ersten Index und `stop` den letzten nicht mehr enthaltenen Index des Abschnitts. Mit `step` können Sie eine Schrittweite angeben, in der die Reihe iteriert werden soll.

Sie müssen nicht alle drei Argumente angeben. Per default ist dann `start=0`, `stop=len(list)` und `step=1`:

## Slicing

Sie können mit der _Slicing_ Syntax auf Teile einer Reihe zugreifen:

```python
slice = list[start:stop:step]
```

Dabei bezeichnet `start` den ersten Index und `stop` den letzten nicht mehr enthaltenen Index des Abschnitts. Mit `step` können Sie eine Schrittweite angeben, in der die Reihe iteriert werden soll.

Sie müssen nicht alle drei Argumente angeben. Per default ist dann `start=0`, `stop=len(list)` und `step=1`:

```python
l[:2] # Der erste Index ist per default start=0
```

```python
l[2:] # Der letzte Index ist per default das Ende der Reihe
```

```python
l[::2] # Jedes zweite Element
```

```python
l[::-1] # Umgekehrte Reihenfolge
```

> **Hinweis:** Slicing ist ein mächtiges Werkzeug um kurzen, prägnanten Code zu schreiben und wird Ihnen bei der Arbeit mit _Numpy Arrays_ noch sehr häufig begegnen.

## Module

Abgesehen von einigen grundlegenden Datentypen und Funktionen wie `print` oder `len` sind viele Funktionen nicht in der Python Standard Library enthalten sondern in **Modulen** ausgelagert. Mit der folgenden Syntax können wir ein Modul **importieren**, um auf die enthaltenen Funktionen zugreifen zu können:

```python
import module
module.function_name()
```

Häufig verwendeten Modulen können wir einen abgekürzten Namen geben:

```python
import module as M
M.function_name()
```

Abweichend von der Standardkonvention ist es empfehlenswert, die Modulnamen in Großbuchstaben zu halten.

Wir können auch nur einzelne Funktionen eines Moduls importieren:

```python
from module import function_name
function_name()
```

## Module bieten vielseitige Funktionalität

Neben eingebauten Modulen wie `math` haben Python-Entwickler eine Vielzahl von Modulen für jeden Anwendungsbereich geschrieben. So können wir mit wenigen Zeilen Code äußerst komplexe Programme schreiben.

> Beispielsweise lesen wir mit `numpy` unseren Datensatz ein, berechnen mit `scipy` einen Fit und plotten beides mit `matplotlib`. Den Umgang mit diesen Modulen lernen wir im nächsten Kapitel.

Funktionen zur Berechnung von Mittelwert und Standardabweichung stellt bspw. `numpy` zur Verfügung:

```python
import numpy as N
li = [1,2,7,3,1,3]
N.mean(li), N.std(li)
```

> Es gibt natürlich nicht nur Module für die wissenschaftliche Anwendung. Python wird höchst vielseitig eingesetzt, sodass Sie bspw. auch
> - eine [Webseite erstellen](http://getpelican.com),
> - einen [Webserver programmieren](http://www.djangoproject.com) oder
> - ein [Spiel entwickeln](http://www.pygame.org) können!

# Numerik mit Numpy

Python Listen sind sehr flexibel, da sie Werte unterschiedlicher Datentypen beinhalten können und einfach verändert werden können (bspw. mit `append`). Diese Flexibilität geht jedoch auf Kosten der Geschwindigkeit, so dass Listen für numerische Berechnungen nicht ideal sind.

Das **Numpy** Modul definiert daher den n-dimensionalen **Array** Datentyp `numpy.ndarray`, der für numerische Berechnungen auf höchst performanten C und Fortran Code zurückgreift.

Arrays können nur Werte eines einzelnen numerischen Datentyps (bspw. floating point Werte) enthalten und sind sehr viel starrer als Listen. Dies ist jedoch für viele wissenschaftliche Anwendung, wie die Arbeit mit Datensätzen, genau was wir brauchen!

Wir importieren das Numpy Modul unter der Abkürzung `N`, da sie prägnanter als die Standardkonvention `np` ist:

```python
import numpy as N # Das Numpy Modul
```

## Numpy Arrays erstellen

Am einfachsten erstellen wir Numpy Arrays aus Python Listen, indem wir die `numpy.array` Funktion verwenden:

```python
a = N.array([ 1, 2, 3, 5, 8, 13 ])
a
```

```python
b = N.array([ [ 1.5, 2.2, 3.1 ], [ 4.0, 5.2, 6.7 ] ])
b
```

Numpy Arrays haben einige **Attribute**, die hilfreiche Informationen über das Array geben:

```python
a.shape, b.shape # Die Länge des Arrays in jeder Dimension
```

```python
a.dtype, b.dtype # Der Datentyp des Arrays
```

> **Erinnerung:** Verwenden Sie die `<TAB>`-Autovervollständigung und die `?`-Dokumentation im Jupyter Notebook wenn Sie nicht wissen, welche Funktionen es gibt oder was diese bewirken!

### Es gibt viele Möglichkeiten, Arrays zu erstellen

- Die `numpy.arange` Funktion erzeugt ein Array mit aufsteigenden Zahlen:

```python
N.arange(10)
```

```python
N.arange(1.5, 2, 0.1)
```

- Außerdem sehr hilfreich sind `numpy.linspace` und `numpy.logspace`, welche eine Anzahl von Werten in linearem oder logarithmischem Abstand zwischen zwei Zahlen generiert:

```python
N.linspace(10, 20, 4)
```

- Wir können mit `numpy.zeros` und `numpy.ones` Arrays erstellen, die mit Nullen oder Einsen gefüllt sind. Indem wir dem Argument `shape` dieser Funktionen statt einem Integer einen Tupel übergeben, können wir auch mehrdimensionale Arrays erzeugen:

```python
N.zeros(5)
```

```python
N.ones((5, 2))
```

## Mit Arrays rechnen

Arrays können mit den Standardoperatoren `+-*/**` **elementweise** kombiniert werden:

```python
x = N.array([1,2,3])
y = N.array([4,5,6])
```

```python
x + 2 * y
```

```python
x ** y
```

> **Achtung:** Für Python-Listen sind diese Operatoren völlig anders definiert!

## Funktionen auf Arrays anwenden

Funktionen aus dem `numpy` Modul sind auf Arrays anwendbar. **Die Funktion wird auf alle Element des Arrays** angewendet und ist typischerweise um einiges schneller als jedes Element einzeln zu berechnen. Grund dafür ist, dass diese Funktionen intern in einer Compilersprache geschrieben und als compilierten Maschinencode ausgeführt werden (wie bei Matlab).

```python
phi = N.linspace(0, 2*N.pi, 10) # 10 Werte zwischen 0 und 2π
N.sin(phi) # Der Sinus jedes dieser Werte
```

Außerdem gibt es viele Funktionen, die Eigenschaften eines Arrays berechnen:

```python
x = N.linspace(0, 10, 100)
N.sum(x), N.mean(x), N.std(x)
```

## Aufgabe 2 - Numpy Arrays

a) Erstelle ein Array `a`, das 11 Werte zwischen $10^{-20}$ und $10^{-10}$ in logarithmischem Abstand enthält.

<!-- grade_id: 201-1a-sol, points: 0, grade: false, solution: true -->
```python
# --- HIER LÖSUNG EINFÜGEN ---
```

<!-- grade_id: 201-1a-test, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_array_equal
try:
    a
except NameError:
    raise NameError("Es gibt keine Variable 'a'. Weise das Array einer Variablen mit diesem Namen zu.")
assert_array_equal(a, [1.00000000e-20, 1.00000000e-19, 1.00000000e-18, 1.00000000e-17, 1.00000000e-16, 1.00000000e-15, 1.00000000e-14, 1.00000000e-13, 1.00000000e-12, 1.00000000e-11, 1.00000000e-10])
print("Jup.")
```

b) Erstelle ein Array `b`, das 2x10 Nullen enthält.

**Hinweis:** Verwende die passende Funktion, die `numpy` bereitstellt.

<!-- grade_id: 201-1b-sol, points: 0, grade: false, solution: true -->
```python
# --- HIER LÖSUNG EINFÜGEN
```

<!-- grade_id: 201-1b-test, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_array_equal
try:
    b
except NameError:
    raise NameError("Es gibt keine Variable 'b'. Weise das Array einer Variablen mit diesem Namen zu.")
assert_array_equal(b, [[0, 0, 0, 0, 0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0, 0, 0, 0, 0]])
print("Gut.")
```

c) Erstelle ein Array `c`, das der Einheitsmatrix in 3 Dimensionen entspricht.

**Hinweis:** Auch hier stellt `numpy` bereits eine passende Funktion bereit.

<!-- grade_id: 201-1c-sol, points: 0, grade: false, solution: true -->
```python
# --- HIER LÖSUNG EINFÜGEN
```

<!-- grade_id: 201-1c-test, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_array_equal
try:
    c
except NameError:
    raise NameError("Es gibt keine Variable 'c'. Weise das Array einer Variablen mit diesem Namen zu.")
assert_array_equal(c, [[1,0,0],[0,1,0],[0,0,1]])
print("Richtig.")
```

d) Erstellen Sie ein Array `d`, das $100$ Werte zwischen $-5$ und $5$ in linearem Abstand enthält. Dies wird eine Raumachse darstellen.

<!-- grade_id: 201-1d-sol, points: 0, grade: false, solution: true -->
```python
# --- HIER LÖSUNG EINFÜGEN
```

<!-- grade_id: 201-1d-test, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_array_almost_equal
try:
    d
except NameError:
    raise NameError("Es gibt keine Variable 'd'. Weise das Array einer Variablen mit diesem Namen zu.")
assert_array_almost_equal(d[:5], [-5., -4.8989899, -4.7979798, -4.6969697, -4.5959596], 4)
print("Stimmt.")
```

### Slicing wählt Teile eines Arrays aus

Die **[Slicing](101%20-%20Grundlagen%20der%20Programmierung%20in%20Python.ipynb#Slicing)** Syntax von Reihen haben wir schon kennengelernt. Sie erlaubt uns, auf einzelne Elemente oder Teile einer Reihe zuzugreifen:

```python
a[start:stop:step]
```

Numpy erweitert diese Syntax auf mehrdimensionale Arrays:

```python
b[start:stop:step, start:stop:step]
```

```python
x = N.arange(10)
```

```python
x[:5]
```

```python
x[::2]
```

Alternativ können wir statt einem Index auch eine **Liste von Indizes** in das Subskript schreiben und erhalten die zugehörigen Elemente aus dem Array:

```python
x = N.array([ 1, 6, 4, 7, 9 ])
indices = [ 1, 0, 2, 1 ]
x[indices]
```

### Slices können auch zugewiesen werden

Wenn ein Slice eines Arrays auf der linken Seite einer Zuweisung steht, wird diesem Teil des Original-Arrays zugewiesen:

```python
x = N.array([ 1, 6, 4, 7, 9 ])
x[x > 4] = 0
x
```

## Aufgabe 3 - Slicing

a) Gegeben ist ein ein Array `x` der Länge `n`. Berechnen Sie das Array `dx` der Länge `n-1` mit den Werten `dx[i] = x[i+1] - x[i]`, ohne die `diff`-Funktion zui benutzen. Verwenden Sie keine Schleifen sondern Slicing!

**Hinweis:** Subtrahieren Sie zwei Arrays, von denen das eine der um 1 versetzte hintere und das andere der vordere Teil von `x` ist.

**Erinnerung:** Mit negativen Zahlen im Subskript wählen Sie Indizes vom Ende einer Reihe aus.

<!-- grade_id: 201-2a-sol, points: 0, grade: false, solution: true -->
```python
x = N.array([ 1, 1, 2, 3, 5, 8 ])
dx = x[1:] - x[:-1] ### --- HIER LÖSUNG EINFÜGEN
```

<!-- grade_id: 201-2a-test, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_array_equal
try:
    dx
except NameError:
    raise NameError("Es gibt keine Variable 'dx'. Weise das Array einer Variablen mit diesem Namen zu.")
assert_array_equal(dx, [0, 1, 1, 2, 3])
print("EZ 😉")
```

# Plots mit Matplotlib

- [Einfaches Plotten](#Einfaches-Plotten)
- [Plots gestalten](#Plots-gestalten)
- [Mehrere Plots in einer Abbildung](#Mehrere-Plots-in-einer-Abbildung)
- [Plots speichern](#Plots-speichern)

Wir können mit Numpy nun numerisch Daten auswerten, und um diese zu visualisieren können wir mit dem **Matplotlib** Modul Plots von wissenschaftlicher Qualität erstellen.

Per Konvention importieren wir das PyPlot Submodul von Matplotlib unter dem Namen `P` (sehr häufig wird `plt` verwendet):

```python
# Zeige Plots direkt im Jupyter Notebook an
%matplotlib inline
# Anschließend können wir das Modul importieren
import matplotlib.pyplot as P
# Numpy brauchen wir immer.
import numpy as N
```

> **Hinweis:** Matplotlib kann mit verschiedenen _Backends_ arbeiten um Plots zu speichern oder auch interaktiv anzuzeigen. Der Aufruf `%matplotlib inline` konfiguriert das Modul für die **statische Anzeige** von Plots im Jupyter Notebook und muss **vor** dem Import von PyPlot ausgeführt werden. Dabei können auch Anzeigeparameter gesetzt werden:
>
> ```python
> %matplotlib inline
> import matplotlib
> # Setzte Anzeigeparameter, z.B.:
> matplotlib.rcParams['figure.figsize'] = (10.0, 8.0)
> import matplotlib.pyplot as plt
> ```
>
> Folgender Aufruf konfiguriert das Jupyter Notebook alternativ für die Verwendung des **interaktiven Backends**:
>
> ```python
> %matplotlib notebook
> ```


## Einfaches Plotten

PyPlot stellt einige grundlegende Funktionen zur Verfügung, mit denen wir Daten schnell plotten können:

- **`plot`** nimmt x- und y-Daten sowie eine Vielzahl von optionalen Argumenten zur Konfiguration an. Fehlen die x-Daten, werden die Indizes der y-Daten verwendet:

```python
P.plot(N.arange(100)**2)
```

Entfernen Sie das '`#`'-Zeichen und schauen Sie sich die Argumente in der Dokumentation an, mit denen wir den Plot konfigurieren können:

```python
#P.plot?
```

Matplotlib bietet unzählige Möglichkeiten, das Aussehen von Plots anzupassen! Ihr könnt oben einige ausprobieren, wie bspw. `color`, `linestyle` und `linewidth`.

- **`scatter`** plottet Datenpunkte statt einer Linie:

```python
P.scatter(N.arange(10), N.arange(10)**2)
```

```python
#P.scatter?
```

## Plots gestalten

Mit Matplotlib können Sie jegliche Plots erstellen, die Sie sich vorstellen können. Überlegen Sie sich zuerst, wie der Plot aussehen soll. Die [Galerie](http://matplotlib.org/gallery.html) kann dabei helfen. Mit einer beherzten Websuche, die meist direkt zur [Dokumentation](http://matplotlib.org/api/pyplot_api.html) führt, finden Sie dann die Funktionen, die Sie benötigen.

> **Hinweis:** Plotten ist eine Kunst.

> **Weiterer Hinweis:** `P.xkcd()`.

### Titel, und Achsen und Legende mit LaTeX

Strings in Matplotlib wie Titel und Achsenbeschriftungen können **LaTeX-Code** enthalten. Text zwischen `$`-Zeichen wird dabei wie gewohnt im _Math-Mode_ gesetzt.

> **Achtung:** Einige Zeichen wie `\t` (`<TAB>`) oder `\n` (`<RETURN>`) werden von Python als Steuerzeichen und daher nicht als LaTeX-Code interpretiert. Markiert Strings, die LaTeX-Code enthalten, daher mit dem Prefix `r` (für _raw_), sodass Steuerzeichen ignoriert werden.

```python
x = N.linspace(0, 2 * N.pi, 100)
# Plot mit Label für Legende
P.plot(x, N.sin(x), label=r'$A \times \sin(\phi)$')
P.plot(x, N.cos(x), label=r'$A \times \cos(\phi)$')
# Titel
P.title('Oszillation')
# Achsenlimits
P.xlim(0, 2 * N.pi)
P.ylim(-1, 1)
# Achsenbeschriftungen
P.xlabel(r'Winkel $\phi \, [\mathrm{rad}]$')
P.ylabel(r'Auslenkung $d \, [\mathrm{cm}]$')
# Legende
P.legend(loc='lower left')
```

## Mehrere Plots in einer Abbildung

Wir können eine Abbildung in ein Gitter mit mehreren Achsen aufteilen. Funktionsaufrufe wie `P.plot` plotten dann immer in die zuletzt erstellte Achse:

```python
# Erstelle eine Abbildung mit Titel
fig = P.figure()
fig.suptitle("Oszillation")

# Erstelle eine Achse in einem Gitter von 2 Zeilen und 1 Spalte
ax_upper = P.subplot(211) # Abkürzung für die 1. Achse im Gitter

# Funktionen plotten nun in diese Achse:
x = N.linspace(0, 2 * N.pi, 100)
P.plot(x, N.sin(x), label=r'$\sin(\phi)$')
P.xlim(0, 2 * N.pi)
P.legend(loc='lower left')

# Erstelle eine zweite Achse im gleichen Gitter, aber an 2. Position
ax_lower = P.subplot(212, sharex=ax_upper, sharey=ax_upper) # Verwende die gleichen x- und y-Limits wie oben

# Funktionen plotten nun in die zweite Achse:
P.plot(x, N.cos(x), label=r'$\cos(\phi)$')
P.xlabel(r'Winkel $\phi \, [\mathrm{rad}]$')
P.legend(loc='lower left')
```

> **Hinweis:** Matplotlib arbeitet mit _Abbildungen_ und _Achsen_. Wenn wir solche nicht explizit erstellen, wie im Beispiel oben, werden sie beim Aufrufen von Funktionen wie `P.plot` automatisch erstellt. Man kann sie auch jederzeit abrufen:
>
> ```python
> fig = P.gcf() # die aktuelle Abbildung
> ax = P.gca() # die aktuelle Achse
> ```
> 
> Wenn Sie genauer bestimmen wollen, auf welche Achse Sie plotten möchten, lönnen Sie sie in Funktionsaufrufen auch angeben oder die Achse direkt verwenden:
>
> ```python
> # Explizit auf Achse `ax` plotten:
> P.plot(x, y, axis=ax) # oder:
> ax.plot(x, y)
> ```

## Plots speichern

Mit `matplotlib.pyplot.savefig` können Sie einen Plot als Bilddatei speichern:

```python
P.savefig('my_plot.png')
```

> **Hinweis:** Um einen Plot im **DIN A4-Format** zu speichern können Sie dessen Größe und Auflösung anpassen:
>
> ```python
> fig = P.gcf()
> fig.set_size_inches(11.69, 8.27)
> P.savefig(filename, dpi=150)
> ```

> **Hinweis:** Um Plots in **LaTeX Dokumente** einzubinden bietet sich das [PGF Vektorformat](https://en.wikipedia.org/wiki/PGF/TikZ) statt Pixelgrafiken wie PNG an. Damit übernimmt LaTeX das Zeichnen des Plots und setzt zudem die enthaltenen Texte selbst, sodass die Fonts und Stile des Dokuments auch im Plot verwendet werden.
>
> Eine Beispiel-Implementierung finden Sie im [TexFig](https://github.com/knly/texfig) Repository.

# Signalverarbeitung in Python

Als erste Schritte in die Signalverarbeitung wollen wir in Python verschiedene Signale erzeugen. 

## Aufgabe 4: Zeitskala

Daten für die Signalverarbeitung liegen im Computer in der Regel als Zahlenfolgen in Arrays vor. Für den Zeitbezug ist eine Zeitskala erforderlich.

Eine Zeitskala ist ein Array, in dem für jedes Element seine Zeitposition notiert ist.

Beispiel:
N.array([0, 0.1, 0.2, 0.3, 0.4, 0.5, 0.6, 0.7, 0.8, 0.9, 1.0]) ist eine Zeitskala von 0 bis 1 Sekunde in 0.1-s-Schritten.

Erzeugen Sie eine Zeitskala in Python für 2 Sekunden (beginnend mit dem Zeitpunkt 0) und speichern Sie sie in die Variable *timeline*. Die Abtastrate betrage 8000 Hz. Plotten Sie die Zeitskala und überzeugen Sie sich von ihrem linearen Verlauf. 

Überlegen Sie sich:
- Wie viele Samples umfasst die Zeitskala? Achten Sie auch auf das Anfang und das Ende der Skala!
- Welchen Definitions- und Wertebereich umfasst sie?

*Hinweis*: Benutzen Sie die Numpy-Funktion *linspace*.

```python
# HIER LÖSUNG EINFÜGEN
```

```python
# Test
assert len(timeline) == 16001
assert timeline[0] == 0.0
assert timeline[-1] == 2.0
print("OK")
```

## Aufgabe 5: Sinussignale
Erstellen Sie ein 2 s langes Sinussignal mit einer Frequenz von 440 Hz und einer Amplitude von 0.7. Benutzen sie die Zeitleiste *timeline* als Basis und speichern Sie den Sinus in die Variable *sine*. Plotten Sie anschließend die ersten 200 Samples des Sinussignals.

Hinweise: Benutzen Sie die von Numpy die Funktion *sin* sowie die Konstante *pi*.

```python
# HIER LÖSUNG EINFÜGEN
```

```python
assert N.max(sine) > 0.699 and N.max(sine) < 0.701
assert sine[0] == 0
assert sine[1] > 0
print("OK!")
```

3b) Bestimmen Sie die Anzahl der lokalen Maxima im Sinussignal und speichern Sie sie in die Variable *num_extrema*.

*Hinweis*: Verwenden Sie die Funktion *scipy.signal.argrelextrema*.

```python
# IHRE LÖSUNG HIER
```

```python
assert num_extrema == 880
print("OK!")
```

# Abtastung, Aliasing und Rekonstruktion

```python
import numpy as N
import scipy.signal as SI
import numpy.fft as FFT
from matplotlib import pylab as P
%matplotlib inline
```

## Aufgabe 6: Aliasing und Frequenzambiguität

Erstellen Sie mit Hilfe der Numpy-Funktion *linspace* zwei Zeitskalen. Beide Skalen haben eine Länge von je einer Sekunde und starten zum Zeitpunkt t=0.
- Die erste Zeitskala habe eine Abtastfrequenz von 2 kHz. Speichern Sie diese Zeitskala in die Variable *timeline2k*.
- Die zweite Zeitskala habe eine Abtastfrequenz von 10 kHz. Speichern Sie diese Zeitskala in die Variable *timeline10k*.

Achten Sie darauf, dass der Zeitpunkt t=1 nicht mehr zur Skala gehört (Benutzen Sie endpoint=False beim Aufruf von *linspace*).

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

```python
assert len(timeline2k) in [2000, 2001]
assert len(timeline10k) in [10000, 10001]
assert timeline2k[0] == 0
assert timeline10k[0] == 0
assert N.abs(timeline2k[-1])-1.0 < 0.001
assert N.abs(timeline2k[-1])-1.0 < 0.00011
print("OK")
```

Erstellen Sie nun auf der Basis von *timeline2k* folgende Signale:
- ein Sinussignal mit der Frequenz von 100 Hz. Speichern Sie es in die Variable *sin100*.
- ein Sinussignal mit der Frequenz von 2100 Hz. Speichern Sie es in die Variable *sin2100*.
- ein Sinussignal mit der Frequenz von 1900 Hz. Speichern Sie es in die Variable *sin1900*.
- ein Cosinussignal mit der Frequenz von 100 Hz. Speichern Sie es in die Variable *cos100*.
- ein Cosinussignal mit der Frequenz von 2100 Hz. Speichern Sie es in die Variable *cos2100*.
- ein Cosinussignal mit der Frequenz von 1900 Hz. Speichern Sie es in die Variable *cos1900*.

Alle Signale haben die Amplitude 1 und die Samplingrate 2000 Hz, da sie auf der Zeitskala *timeline2k* basieren.

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

```python
for sine in [sin100, sin2100, sin1900]:
    assert sine[0] == 0
    assert len(SI.argrelextrema(sine, N.greater)[0]) == 100
for cosine in [cos100, cos2100, cos1900]:
    assert cosine[0] == 1
    assert len(SI.argrelextrema(cosine, N.greater)[0])  == 99
print("OK!")
```

Plotten Sie die ersten 100 Samples von den Signalen *sin100* und *sin2100*. Was fällt auf?

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

Plotten Sie die ersten 100 Samples von den Signalen *sin100* und *sin1900*. Was fällt auf?

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

Plotten Sie die ersten 100 Samples von den Signalen *cos100* und *cos2100*. Was fällt auf?

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

Plotten Sie die ersten 100 Samples von den Signalen *cos100* und *cos1900*. Vergleichen Sie das Ergebnis mit dem Versuch mit *sin1900*. Was fällt auf?

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

Erstellen Sie nun zwei je 1 s lange Sinussignale mit den Frequenzen 100 Hz und 2100 Hz, aber diesmal mit der Abtastrate 10 kHz. Benutzen Sie die Zeitleiste *timeline10k*. Speichern Sie die Signale in die Variablen *extsin100* und *extsin2100*.

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

```python
assert extsin100[0] == 0
assert extsin2100[0] == 0
assert len(SI.argrelextrema(extsin100, N.greater)[0]) == 100
assert len(SI.argrelextrema(extsin2100, N.greater)[0]) == 2100
print("OK")
```

Plotten Sie die ersten 100 Samples von den Signalen *extsin100* und *extsin2100* und vergleichen Sie das Ergebnis mit dem Plot von *sin100* und *sin2100*. Was fällt auf? Wie erklären Sie sich dieses Ergebnis?

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

## Aufgabe 7: Die sinc-Funktion

Lassen Sie sich über die Fragezeichen-Funktion von Jupyter die Dokumentation der Numpy-Funktion *sinc* anzeigen und beantworten Sie folgende Fragen:
- Wie lautet die mathematische Definition der sinc-Funktion?
- Was ist der Unterschied zwischen sinc und si?
- Wie ist sinc(0) definiert?

Für die si-Funktion gilt:
$$ si(x) = \frac{\sin(x)}{x}$$

```python
N.sinc?
```

Erstellen Sie nun eine Zeitleiste von -6 bis +6 mit 1301 Punkten. Speichern Sie die Zeitleiste in die Variable *timeline_sinc*. Der Zeitpunkt $t=6$ soll in der Zeitleiste enthalten sein.

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

```python
assert timeline_sinc[0] == -6
assert len(timeline_sinc) == 1301
assert timeline_sinc[-1] == 6
print("OK!")
```

Wenden Sie nun die *sinc*-Funktion auf *timeline_sinc* an und speichern Sie das Ergebnis in die Variable *sinc*. Plotten Sie das Ergebnis gegen die Zeitleiste.

```python
# HIER IHRE LÖSUNG EINFÜGEN
```

Finden Sie die Nullstellen von *sinc* heraus.
- Als Nullstelle gelten alle Stellen, bei denen ein Vorzeichenwechsel passiert. Dies ist wichtig, weil man sich bei digitalen Signalen nicht darauf verlassen kann, dass pünktlich auf einer Nullstelle das Signal abgetastet wird.
- Benutzen Sie die Numpy-Funktionen *sign*, *diff* und *nonzero*, um die Samplepositionen von Vorzeichenwechseln zu ermitteln.
- Geben Sie die ermittelten Positionen in Samples sowie in Zeitpunkten auf der Timeline aus.


```python
samplePositions = N.nonzero(N.diff(N.sign(sinc)))
timePositions = timeline_sinc[samplePositions]
print (samplePositions, timePositions)
```

## Anhang: Python installieren

Eine einfache Möglichkeit, den Python Interpreter, Jupyter und wissenschaftliche Packages auf Mac OS X, Linux und Windows zu installieren, ist die [Anaconda](https://www.continuum.io/downloads) Distribution:

1. Installieren Sie [Anaconda](https://www.continuum.io/downloads).
2. Öffnen Sie den [Anaconda Navigator](https://docs.continuum.io/anaconda/navigator/) und starten Sie die _Jupyter Notebook_ App.
3. Im Browserfenster können Sie nun durch euer Dateisystem navigieren, Jupyter Notebooks erstellen und darin Python-Code ausführen.

> Wenn Sie Erfahrung im Umgang mit der Kommandozeile und Package Managern habt können Sie den Python Interpreter auch anderweitig installieren, unter Mac OS X bspw. mit [Homebrew](http://brew.sh). Packages wie `numpy`, `matplotlib` und `scipy` installieren Sie dann mit dem Python Package Manager `pip`. Auch das Jupyter Notebook installiert ihr mit `pip install jupyter`. In der Anaconda Distribution sind diese Packages bereits enthalten.

> Haben Sie `jupyter` über die Anaconda Distribution oder `pip` installiert, können SIe statt den Anaconda Navigator zu verwenden auch mit der Kommandozeile in ein Verzeichnis navigieren und den Befehl
>
> ```bash
> jupyter notebook
> ```
>
> ausführen. Es öffnet sich ein Browserfenster, in dem Sie Notebooks verwalten, bearbeiten und ausführen können.

> Wenn Sie nicht die Anaconda Distribution verwenden, stellen Sie sicher, dass die benötigten Packages für diesen Kurs installiert sind. Sie sind in der `requirements.txt` Datei aufgelistet und können einfach mit `pip` installiert werden:
>
> ```bash
> pip install -r requirements.txt
> ```

## Python Code in der Shell und in Skripten ausführen

Sie haben eine funktionsfähige Python Installation, wenn Sie in der Kommandozeile den Befehl `python --version` ausführen und eine Versionsnummer erhalten. Nun haben S neben Jupyter Notebooks viele weitere Möglichkeiten, Python-Code auszuführen:

Um eure Python Installation zu testen könnt ihr die _Python Shell_ öffnen, indem ihr in der Kommandozeile einfach `python` ausführt. Hier können SIe Python Code schreiben und direkt ausführen. Versuchen Sie es einmal mit `print("Hello World!")`. In der Shell können Sie testen, ob Packages wie `numpy`, `matplotlib` und `scipy` installiert sind, indem Sie versuchen, sie mit `import` zu importieren. Die Shell wird mit `exit()` verlassen.

Natürlich eignet sich die Shell nicht für ernsthaftes Programmieren. Für Skripte oder größere Programme, die Aufgaben ausführen sollen, schreiben wir den Code daher in `.py`-Textdateien. Ein guter Texteditor wie [GitHub's Atom](https://atom.io), Microsoft's [Visual Studio Code](https://code.visualstudio.com) oder sogar eine vollwertige Entwicklungsumgebung wie [PyCharm](https://www.jetbrains.com/pycharm/) ist dabei essentiell. Eine Alternative ist das an MATLAB angelehnte [Spyder](https://pythonhosted.org/spyder/).

Erstellen Sie eine Datei `helloworld.py` mit dem Inhalt `print("Hello World!")`. Nun könneb Sie das Skript in der Kommandozeile ausführen:

    python path/to/helloworld.py

## Referenzmaterialien

### Offizielle Python-Dokumentation
- https://docs.python.org
- Standardbibliothek: https://docs.python.org/library/3/index.html

### Deutschsprachiges Python-Intensivtutorial
- https://python-kurs.eu

### Python für MATLAB-Umsteiger
- https://numpy.org/doc/stable/user/numpy-for-matlab-users.html
- https://realpython.com/matlab-vs-python/
